```python
version = "REPLACE_PACKAGE_VERSION"
```

# Data Manipulation
## School of Information, University of Michigan

## Assignment 3
All questions are weighted the same in this assignment. This assignment requires more individual learning then the last one did - you are encouraged to check out the [pandas documentation](http://pandas.pydata.org/pandas-docs/stable/) to find functions or methods you might not have used yet, or ask questions on [Stack Overflow](http://stackoverflow.com/) and tag them as pandas and python related. All questions are worth the same number of points except question 1 which is worth 20% of the assignment grade.

**Note**: Questions 2-13 rely on your question 1 answer.

```python
import pandas as pd
import numpy as np

# Filter all warnings. If you would like to see the warnings, please comment the two lines below.
import warnings
warnings.filterwarnings('ignore')
```

### Question 1
Load the energy data from the file `assets/Energy Indicators.xls`, which is a list of indicators of [energy supply and renewable electricity production](assets/Energy%20Indicators.xls) from the [United Nations](http://unstats.un.org/unsd/environment/excel_file_tables/2013/Energy%20Indicators.xls) for the year 2013, and should be put into a DataFrame with the variable name of **Energy**.

Keep in mind that this is an Excel file, and not a comma separated values file. Also, make sure to exclude the footer and header information from the datafile. The first two columns are unneccessary, so you should get rid of them, and you should change the column labels so that the columns are:

`['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]`

Convert `Energy Supply` to gigajoules (there are 1,000,000 gigajoules in a petajoule). For all countries which have missing data (e.g. data with "...") make sure this is reflected as `np.NaN` values.

Rename the following list of countries (for use in later questions):

```"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"```

There are also several countries with parenthesis in their name. Be sure to remove these, e.g. `'Bolivia (Plurinational State of)'` should be `'Bolivia'`.

Next, load the GDP data from the file `assets/world_bank.csv`, which is a csv containing countries' GDP from 1960 to 2015 from [World Bank](http://data.worldbank.org/indicator/NY.GDP.MKTP.CD). Call this DataFrame **GDP**. 

Make sure to skip the header, and rename the following list of countries:

```"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"```

Finally, load the [Sciamgo Journal and Country Rank data for Energy Engineering and Power Technology](http://www.scimagojr.com/countryrank.php?category=2102) from the file `assets/scimagojr-3.xlsx`, which ranks countries based on their journal contributions in the aforementioned area. Call this DataFrame **ScimEn**.

Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15). 

The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
       'Citations per document', 'H index', 'Energy Supply',
       'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008',
       '2009', '2010', '2011', '2012', '2013', '2014', '2015'].

*This function should return a DataFrame with 20 columns and 15 entries.*

<!-- grade_id: cell-bce4d6f2ecdd1297, points: 0, grade: false, solution: true -->
```python
def answer_one():
    ### BEGIN SOLUTION
    # Competency: reading files in pandas, df manipulation, regex, joining datasets
    Energy = pd.read_excel('assets/Energy Indicators.xls', skiprows=17, skip_footer=38)
    Energy.columns = [0, 1, 'Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable']
    Energy = Energy.drop([0,1],1)
    Energy['Country'] = Energy['Country'].replace(regex=True, to_replace='\d', value='')
    Energy.set_index('Country', inplace=True)

    Energy.rename(index={"Republic of Korea":"South Korea", 
                         "United States of America":"United States",
                         "United Kingdom of Great Britain and Northern Ireland":"United Kingdom",
                         "China, Hong Kong Special Administrative Region":"Hong Kong"}, inplace=True)

    Energy = Energy.replace('...', np.nan)
    Energy['Energy Supply'] *= 10**6
    Energy.index = Energy.index.str.replace(r" \(.*\)","")

    GDP = pd.read_csv('assets/world_bank.csv', skiprows=4, index_col=0)
    GDP.rename(index={"Korea, Rep.":"South Korea", 
                      "Iran, Islamic Rep.": "Iran",
                      "Hong Kong SAR, China": "Hong Kong"}, inplace=True)

    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx', index_col=1)
    
    df = ScimEn.join([Energy, GDP.ix[:,-10:]], how='inner')
    return df.sort_values('Rank').head(15)
    ### END SOLUTION
```

<!-- grade_id: cell-780b5a4da845dbc3, points: 5, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
assert (answer_one().columns == ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
                                 'Citations per document', 'H index', 'Energy Supply','Energy Supply per Capita', '% Renewable', 
                                 '2006', '2007', '2008','2009', '2010', '2011', '2012', '2013', '2014', '2015']).all() , "Please name the columns correctly!"

### END HIDDEN TESTS
```

<!-- grade_id: cell-2e9e8d2d1a53e343, points: 5, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
assert answer_one().shape == (15,20), "Your function should return a DataFrame with 20 columns and 15 entries."
### END HIDDEN TESTS
```

<!-- grade_id: cell-74b5f0b971379f64, points: 10, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
assert (answer_one().index == ['China', 'United States', 'Japan', 'United Kingdom',
       'Russian Federation', 'Canada', 'Germany', 'India', 'France',
       'South Korea', 'Italy', 'Spain', 'Iran', 'Australia', 'Brazil']).all(), "Wrong Countries!"

assert list(answer_one().Rank) == [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15], "Wrong orders of Rank!"

assert list(answer_one()['Energy Supply']) == [127191000000.0,90838000000.0,18984000000.0,7920000000.0,30709000000.0,
                                          10431000000.0,13261000000.0,33195000000.0,10597000000.0,11007000000.0,
                                          6530000000.0,4923000000.0,9172000000.0,5386000000.0,12149000000.0], "Wrong Energy Supply Data!"

assert list(answer_one()['2011']) == [6612490294273.65,15204019634600.0,5473738114222.05,2450911011124.85,
                                      1589942641764.29,1664086754804.89,3542371125827.81,1821872415623.62,
                                      2702031788079.47,1134795571930.64,2137439133774.83,1417354966887.42,
                                      485330901134.448,1169431057881.91,2295244758549.53], "Wrong GDP Data!"

### END HIDDEN TESTS
```

### Question 2
The previous question joined three datasets then reduced this to just the top 15 entries. When you joined the datasets, but before you reduced this to the top 15 items, how many entries did you lose?

*This function should return a single number.*

```python
%%HTML
<svg width="800" height="300">
  <circle cx="150" cy="180" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="blue" />
  <circle cx="200" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="red" />
  <circle cx="100" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="green" />
  <line x1="150" y1="125" x2="300" y2="150" stroke="black" stroke-width="2" fill="black" stroke-dasharray="5,3"/>
  <text x="300" y="165" font-family="Verdana" font-size="35">Everything but this!</text>
</svg>
```

<!-- grade_id: cell-b0c3202c000aada4, points: 0, grade: false, solution: true -->
```python
def answer_two():
    ### BEGIN SOLUTION
    # Competency: joining datasets, sets
    import pandas as pd
    import numpy as np

    Energy = pd.read_excel('assets/Energy Indicators.xls', skiprows=17, skip_footer=38)
    Energy.columns = [0, 1, 'Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable']
    Energy = Energy.drop([0,1],1)
    Energy['Country'] = Energy['Country'].replace(regex=True, to_replace='\d', value='')
    Energy.set_index('Country', inplace=True)

    Energy.rename(index={"Republic of Korea":"South Korea", 
                         "United States of America":"United States",
                         "United Kingdom of Great Britain and Northern Ireland":"United Kingdom",
                         "China, Hong Kong Special Administrative Region":"Hong Kong"}, inplace=True)

    Energy = Energy.replace('...', np.nan)
    Energy['Energy Supply'] *= 10**6
    Energy.index = Energy.index.str.replace(r" \(.*\)","")

    GDP = pd.read_csv('assets/world_bank.csv', skiprows=4, index_col=0)
    GDP.rename(index={"Korea, Rep.":"South Korea", 
                      "Iran, Islamic Rep.": "Iran",
                      "Hong Kong SAR, China": "Hong Kong"}, inplace=True)

    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx', index_col=1)
    return len((set(GDP.index)|set(Energy.index)|set(ScimEn.index)) - (set(GDP.index)&set(Energy.index)&set(ScimEn.index)))
    ### END SOLUTION
```

<!-- grade_id: cell-be24cfcaa87ab071, points: 6.66, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
assert type(answer_two()) == int, "Please return an int number."

assert answer_two() == 156, "Wrong number!"
### END HIDDEN TESTS
```

### Question 3
What are the top 15 countries for average GDP over the last 10 years?

*This function should return a Series named `avgGDP` with 15 countries and their average GDP sorted in descending order.*

<!-- grade_id: cell-8c3d74335c0d489a, points: 0, grade: false, solution: true -->
```python
def answer_three():
    ### BEGIN SOLUTION
    # Competency: indexing, math fn, sorting
    Top15 = answer_one()
    years = [str(x) for x in range(2006,2016)]
    Top15['avgGDP'] = Top15[years].mean(1)
    return Top15['avgGDP'].sort_values(ascending=False)
    ### END SOLUTION
```

<!-- grade_id: cell-aaaa11ef7d26f4cf, points: 6.66, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
assert list(answer_three().index) == ['United States', 'China', 'Japan', 'Germany', 'France',
       'United Kingdom', 'Brazil', 'Italy', 'India', 'Canada',
       'Russian Federation', 'Spain', 'Australia', 'South Korea', 'Iran'],"Wrong Countries!"


assert list(answer_three()) == [15364344302990.0,6348608932836.1,5542207638235.176,3493025339072.8477,2681724635761.589,
 2487906661418.4175,2189794143774.905,2120175089933.7761,1769297396603.86,1660647466307.512,
 1565459478480.661,1418078278145.694,1164042729991.427,1106714508244.852,444155754051.095], "Wrong GDP Data!"
### END HIDDEN TESTS
```

### Question 4
By how much had the GDP changed over the 10 year span for the country with the 6th largest average GDP?

*This function should return a single number.*

<!-- grade_id: cell-7f77d099e3e0bbee, points: 0, grade: false, solution: true -->
```python
def answer_four():
    ### BEGIN SOLUTION
    # Competency: indexing, broadcasting
    Top15 = answer_one()
    return Top15.ix['United Kingdom', '2015'] - Top15.ix['United Kingdom', '2006']
    ### END SOLUTION
```

<!-- grade_id: cell-564dd8e5e24b0f83, points: 6.66, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
assert round(answer_four(),1) == 246702696075.4, "Wrong country or number!"
### END HIDDEN TESTS
```

### Question 5
What is the mean energy supply per capita?

*This function should return a single number.*

<!-- grade_id: cell-58e79d558e982eef, points: 0, grade: false, solution: true -->
```python
def answer_five():
    ### BEGIN SOLUTION
    # Competency: math fn
    Top15 = answer_one()
    return Top15['Energy Supply per Capita'].mean()
    ### END SOLUTION
```

<!-- grade_id: cell-30cc66180851638c, points: 6.66, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
assert round(answer_five(),1) == 157.6, "Wrong number!"
### END HIDDEN TESTS
```

### Question 6
What country has the maximum % Renewable and what is the percentage?

*This function should return a tuple with the name of the country and the percentage.*

<!-- grade_id: cell-b6824b78e74619f9, points: 0, grade: false, solution: true -->
```python
def answer_six():
    ### BEGIN SOLUTION
    # Competency: math fn
    Top15 = answer_one()
    return Top15['% Renewable'].idxmax(), Top15['% Renewable'].max()
    ### END SOLUTION
```

<!-- grade_id: cell-2bd201c5c7bdd80f, points: 6.66, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
assert type(answer_six()) ==tuple, "Your function should return a tuple!"

assert answer_six()[0]=='Brazil', "Wrong Country!"
assert round(answer_six()[1]) ==70, "Wrong percentage!"

### END HIDDEN TESTS
```

### Question 7
Create a new column that is the ratio of Self-Citations to Total Citations. 
What is the maximum value for this new column, and what country has the highest ratio?

*This function should return a tuple with the name of the country and the ratio.*

<!-- grade_id: cell-a4f39737f38aa53c, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    ### BEGIN SOLUTION
    # Competency: math fn, broadcasting
    Top15 = answer_one()
    Top15['Self_to_Cite'] = Top15['Self-citations'] / Top15['Citations']
    return Top15['Self_to_Cite'].idxmax(), Top15['Self_to_Cite'].max()
    ### END SOLUTION
```

<!-- grade_id: cell-b7a163e9231b88c9, points: 6.66, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
assert type(answer_seven()) ==tuple, "Your function should return a tuple!"

assert answer_seven()[0]=='China', "Wrong Country!"
assert round(answer_seven()[1],1) ==0.7, "Wrong percentage!"

### END HIDDEN TESTS
```

### Question 8

Create a column that estimates the population using Energy Supply and Energy Supply per capita. 
What is the third most populous country according to this estimate?

*This function should return the name of the country*

<!-- grade_id: cell-9ca58137846b84d6, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    ### BEGIN SOLUTION
    # Competency: Broadcasting, sorting
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    return Top15['PopEst'].sort_values(ascending=False).index[2]
    ### END SOLUTION
```

<!-- grade_id: cell-3f3620c88df08b20, points: 0, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
assert type(answer_eight()) ==str, "Your function should return a string!"
assert answer_eight()=='United States', "Wrong Country!"
### END HIDDEN TESTS
```

### Question 9
Create a column that estimates the number of citable documents per person. 
What is the correlation between the number of citable documents per capita and the energy supply per capita? Use the `.corr()` method, (Pearson's correlation).

*This function should return a single number.*

*(Optional: Use the built-in function `plot9()` to visualize the relationship between Energy Supply per Capita vs. Citable docs per Capita)*

<!-- grade_id: cell-033679ea456bfb9d, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    ### BEGIN SOLUTION
    # Competency: Broadcasting, math fn, correlation, visualization
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    return Top15['Citable docs per Capita'].corr(Top15['Energy Supply per Capita'])
    ### END SOLUTION
```

```python
def plot9():
    import matplotlib as plt
    %matplotlib inline
    
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15.plot(x='Citable docs per Capita', y='Energy Supply per Capita', kind='scatter', xlim=[0, 0.0006])
```

<!-- grade_id: cell-3cb5c699065a4a20, points: 6.66, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
assert round(answer_nine(),1)==0.8, "Wrong number!"
### END HIDDEN TESTS
```

### Question 10
Create a new column with a 1 if the country's % Renewable value is at or above the median for all countries in the top 15, and a 0 if the country's % Renewable value is below the median.

*This function should return a series named `HighRenew` whose index is the country name sorted in ascending order of rank.*

<!-- grade_id: cell-0fdf60e64bf1a4f9, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    ### BEGIN SOLUTION
    # Competency: df querying, math fn, variable encoding
    Top15 = answer_one()
    Top15['HighRenew'] = np.where(Top15['% Renewable'] >= Top15['% Renewable'].median(), 1, 0)
    return Top15['HighRenew']
    ### END SOLUTION
```

<!-- grade_id: cell-b29a631fd9a7730f, points: 6.66, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
assert list(answer_ten().index)==['China', 'United States', 'Japan', 'United Kingdom',
       'Russian Federation', 'Canada', 'Germany', 'India', 'France',
       'South Korea', 'Italy', 'Spain', 'Iran', 'Australia', 'Brazil'], "Wrong country order！"

assert list(answer_ten()) == [1, 0, 0, 0, 1, 1, 1, 0, 1, 0, 1, 1, 0, 0, 1], "Check the values in your new column."

### END HIDDEN TESTS
```

### Question 11
Use the following dictionary to group the Countries by Continent, then create a DataFrame that displays the sample size (the number of countries in each continent bin), and the sum, mean, and std deviation for the estimated population of each country.

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
```

*This function should return a DataFrame with index named Continent `['Asia', 'Australia', 'Europe', 'North America', 'South America']` and columns `['size', 'sum', 'mean', 'std']`*

<!-- grade_id: cell-a5e0c0df27304f98, points: 0, grade: false, solution: true -->
```python
# Competency: mapping, groupby, agg
def answer_eleven():
    ### BEGIN SOLUTION
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    ContinentDict  = {'China':'Asia', 
                      'United States':'North America', 
                      'Japan':'Asia', 
                      'United Kingdom':'Europe', 
                      'Russian Federation':'Europe', 
                      'Canada':'North America', 
                      'Germany':'Europe', 
                      'India':'Asia',
                      'France':'Europe', 
                      'South Korea':'Asia', 
                      'Italy':'Europe', 
                      'Spain':'Europe', 
                      'Iran':'Asia',
                      'Australia':'Australia', 
                      'Brazil':'South America'}

    Top15['Continent'] = pd.Series(ContinentDict)

    grouped = Top15.groupby(['Continent'])
    grouped.size()

    return grouped['PopEst'].agg([np.size, np.sum, np.mean, np.std])
    ### END SOLUTION
```

<!-- grade_id: cell-18d1a07971b25743, points: 6.66, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
res = answer_eleven()
assert list(res.index)==['Asia', 'Australia', 'Europe', 'North America', 'South America'], "Wrong index！"
assert list(res.columns) == ['size', 'sum', 'mean', 'std'], "Wrong column names!"
assert list(res['size']) == [5.0, 1.0, 6.0, 2.0, 1.0], "Check the size column."
assert round(res.ix['Asia','sum'],1) == 2898666386.6, "Check the sum column."
assert round(res.ix['South America','mean'],1) == 205915254.2, "Check the mean column."
assert round(res.ix['Asia','std'],1) == 679097888.4, "Check the std column."
assert np.isnan(answer_eleven().ix['Australia','std']) == True, "Check the std column."

### END HIDDEN TESTS
```

### Question 12
Cut % Renewable into 5 bins. Group Top15 by the Continent, as well as these new % Renewable bins. How many countries are in each of these groups?

*This function should return a Series with a MultiIndex of `Continent`, then the bins for `% Renewable`. Do not include groups with no countries.*

<!-- grade_id: cell-2ecd9a4076abd8f0, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    ### BEGIN SOLUTION
    # Competency: cut, groupby, math fn
    Top15 = answer_one()
    ContinentDict  = {'China':'Asia', 
                      'United States':'North America', 
                      'Japan':'Asia', 
                      'United Kingdom':'Europe', 
                      'Russian Federation':'Europe', 
                      'Canada':'North America', 
                      'Germany':'Europe', 
                      'India':'Asia',
                      'France':'Europe', 
                      'South Korea':'Asia', 
                      'Italy':'Europe', 
                      'Spain':'Europe', 
                      'Iran':'Asia',
                      'Australia':'Australia', 
                      'Brazil':'South America'}

    Top15['Continent'] = pd.Series(ContinentDict)    
    return Top15.groupby([Top15['Continent'], pd.cut(Top15['% Renewable'], 5)]).size()
    ### END SOLUTION

```

<!-- grade_id: cell-6c665602d6babab9, points: 6.66, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS
assert list(answer_twelve())==[4, 1, 1, 1, 3, 2, 1, 1, 1], "Check the number of countries in each group."
assert list(answer_twelve().index.names) == ['Continent', '% Renewable'], "Check the MultiIndex of your Dataframe."
### END HIDDEN TESTS
```

### Question 13
Convert the Population Estimate series to a string with thousands separator (using commas). Use all significant digits (do not round the results).

e.g. 12345678.90 -> 12,345,678.90

*This function should return a series `PopEst` whose index is the country name and whose values are the population estimate string*

<!-- grade_id: cell-58eb0ee0921d93fb, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    ### BEGIN SOLUTION
    # Competency: lambdas, format
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    return Top15['PopEst'].apply(lambda x: "{:,}".format(x))
    ### END SOLUTION
```

<!-- grade_id: cell-10fee7228cf973f6, points: 6.74, grade: true, solution: false -->
```python
# Hidden Tests
### BEGIN HIDDEN TESTS

assert list(answer_thirteen().index) == ['China', 'United States', 'Japan', 'United Kingdom',
                                         'Russian Federation', 'Canada', 'Germany', 'India', 'France',
                                         'South Korea', 'Italy', 'Spain', 'Iran', 'Australia', 'Brazil'], "Wrong index!"
assert list(answer_thirteen())==['1,367,645,161.2903225','317,615,384.61538464','127,409,395.97315437','63,870,967.741935484','143,500,000.0',
 '35,239,864.86486486','80,369,696.96969697','1,276,730,769.2307692','63,837,349.39759036','49,805,429.864253394',
 '59,908,256.880733944','46,443,396.2264151','77,075,630.25210084','23,316,017.316017315','205,915,254.23728815'], "Check the values of your population estimate strings."


### END HIDDEN TESTS
```

### Optional

Use the built in function `plot_optional()` to see an example visualization.

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")
```