```python
version = "REPLACE_PACKAGE_VERSION"
```

For this problem set, we'll be using the Jupyter notebook:

![](assets/jupyter.png)

### Part A: Euclidean distance (3 points)

Recall that the Euclidean distance $d$ is given by the following equation:

$$
d(a, b) = \sqrt{\sum_{i=1}^N (a_i - b_i) ^ 2}
$$

In NumPy, this is a fairly simple computation because we can rely on array computations and the `np.sum` function to do all the heavy lifting for us.

<div class="alert alert-success">
Complete the function <code>euclidean_distance</code> below to compute $d(a,b)$, as given by the equation above. Note that you can compute the square root using <code>np.sqrt</code>.
</div>

```python
import numpy as np
```

<!-- grade_id: euclidean_distance, points: 0, grade: false, solution: true -->
```python
def euclidean_distance(a, b):
    """Computes the Euclidean distance between a and b.
    
    Hint: your solution can be done in a single line of code!
    
    Parameters
    ----------
    a, b : numpy arrays or scalars with the same size
    
    Returns
    -------
    the Euclidean distance between a and b
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<div class="alert alert-warning">Remember that you need to execute the cell above (with your definition of <code>euclidean_distance</code>), and then run the cell below to check your answer. If you make changes to the cell with your answer, you will need to <i>first</i> re-run that cell, and <i>then</i> re-run the test cell to check your answer again.</div>

<!-- grade_id: cell-25477b5440736077, points: 1, grade: true, solution: false -->
```python
a = np.array([1, 2, 3])
b = np.array([4, 5, 6])
assert type(euclidean_distance(a, b)) == np.float64
```

<!-- grade_id: test_euclidean_distance, points: 2, grade: true, solution: false -->
```python
# Hidden Tests
```