Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python

```

---

For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## Part A (2 points)

Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    if n < 1:
        raise ValueError
    s = []
    for i in range(n):
        s.append(i**2)
    return s
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

```python
squares(10)
```

Output:
```
[0, 1, 4, 9, 16, 25, 36, 49, 64, 81]
```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-3-f3fef5b9ed4e> in <module>
      1 """Check that squares returns the correct output for several inputs"""
----> 2 assert squares(1) == [1]
      3 assert squares(2) == [1, 4]
      4 assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
      5 assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]

AssertionError:
```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
try:
    squares(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    squares(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

---

## Part B (3 point)

Using your `squares` function, write a function that computes the sum of the squares of the numbers from 1 to $n$. Your function should call the `squares` function -- it should NOT reimplement its functionality.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """Compute the sum of the squares of numbers from 1 to n."""
    total = 0
    s = squares(n)
    for i in range(len(s)):
        total += s[i]
    return total
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```python
sum_of_squares(10)
```

Output:
```
285
```

<!-- grade_id: correct_sum_of_squares_one, points: 1, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
assert sum_of_squares(1) == 1
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-7-d1f5f7b50548> in <module>
      1 """Check that sum_of_squares returns the correct answer for various inputs."""
----> 2 assert sum_of_squares(1) == 1

AssertionError:
```

<!-- grade_id: correct_sum_of_squares, points: 1, grade: true, solution: false -->
```python
assert sum_of_squares(2) == 5
### BEGIN HIDDEN TESTS
assert sum_of_squares(10) == 385
assert sum_of_squares(11) == 506
### END HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-8-9937fff5e587> in <module>
----> 1 assert sum_of_squares(2) == 5
      2 ### BEGIN HIDDEN TESTS
      3 assert sum_of_squares(10) == 385
      4 assert sum_of_squares(11) == 506
      5 ### END HIDDEN TESTS

AssertionError:
```

<!-- grade_id: sum_of_squares_uses_squares, points: 1, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on squares."""
orig_squares = squares
del squares
try:
    sum_of_squares(1)
except NameError:
    pass
else:
    raise AssertionError("sum_of_squares does not use squares")
finally:
    squares = orig_squares
```